Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

NAME = "Oliver Collins" <br>
COLLABORATORS = "Rahul Depa"

---

```scala
// import $file.hw6stdlib
// import hw6stdlib._


// Nats
sealed trait Nat
case object Zero extends Nat
case class Succ(pred : Nat) extends Nat

val one = Succ(Zero)
val two = Succ(one)
val three = Succ(two)
val four = Succ(three)
val five = Succ(four)
val six = Succ(five)
val seven = Succ(six)
val eight = Succ(seven)
val nine = Succ(eight)
val ten = Succ(nine)

def nat_plus(x : Nat, y : Nat) : Nat = x match {
    case Zero    => y
    case Succ(x) => Succ(nat_plus(x, y))
}

def nat_to_int(x : Nat) : Int = x match {
    case Zero => 0
    case Succ( x ) => 1 + nat_to_int(x)
}

def nat_to_str(x : Nat) = nat_to_int(x).toString()

def print_nat(x : Nat) = println(nat_to_str(x))

def nat_mult(x : Nat, y : Nat) : Nat = x match {
    case Zero    => Zero
    case Succ(x) => nat_plus(nat_mult(x,y), y)
}

def nat_minus(x : Nat, y : Nat) : Nat = (x, y) match {
    case (Zero, _)          => Zero
    case (x, Zero)          => x
    case (Succ(x), Succ(y)) => nat_minus(x, y)
}

def nat_pow(x : Nat, y : Nat) : Nat = x match {
    case Zero       => one
    case Succ(Zero) => y 
    case Succ(x)    => nat_mult(y, nat_pow(x, y))
}

def nat_lte(x : Nat, y : Nat) : Bool = (x, y) match {
    case (Zero, y)          => True
    case (x, Zero)          => False
    case (Succ(x), Succ(y)) => nat_lte(x, y)
}

def nat_eq(x : Nat, y : Nat) : Bool = (x, y) match {
    case (Zero, Zero)       => True
    case (Succ(x), Succ(y)) => nat_eq(x, y)
    case _                  => False
}

// Integers
sealed trait Integer
case class Positive(x : Nat) extends Integer
case class Negative(x : Nat) extends Integer

def int_to_str(x : Integer) : String = x match {
    case Positive(x) => nat_to_str(x)
    case Negative(x) => "-" + nat_to_str(x)
}

def print_integer(x : Integer) = println(int_to_str(x))

def abs(x : Integer) : Nat = x match {
    case Positive(x) => x
    case Negative(x) => x
}

def negate(x : Integer) : Integer = x match {
    case Positive(x) => Negative(x)
    case Negative(x) => Positive(x)
}

def plus(n : Integer, m : Integer) : Integer = (n, m) match {
    case (Positive(x), Positive(y)) => Positive(nat_plus(x, y))
    case (Negative(x), Negative(y)) => Negative(nat_plus(x,y))
    case (Negative(x), Positive(y)) => nat_lte(x, y) match {
        case True  => Positive(nat_minus(y, x))
        case False => Negative(nat_minus(x, y))
    }
    case (Positive(x), Negative(y)) => nat_lte(x, y) match {
        case True  => Negative(nat_minus(y, x))
        case False => Positive(nat_minus(x, y))
    }    
}

def minus(x : Integer, y : Integer) : Integer = plus(x, negate(y))

def mult(x : Integer, y : Integer) : Integer = (x,y) match {
    case (Positive(x), Positive(y)) => Positive(nat_mult(x, y))   
    case (Negative(x), Negative(y)) => Positive(nat_mult(x, y))
    case (Negative(x), Positive(y)) => Negative(nat_mult(x, y))
    case (Positive(x), Negative(y)) => Negative(nat_mult(x, y))
}

def pow(x : Integer, y : Nat) : Integer = y match {
    case Zero => Positive(Succ(Zero))
    case Succ(y) => mult(x, pow(x, y))
}

def int_eq(x : Integer, y : Integer) : Bool = (x, y) match {
    case (Positive(x), Positive(y)) => nat_eq(x, y)
    case (Negative(x), Negative(y)) => nat_eq(x, y)
    case _                          => False
}

// Booleans
sealed trait Bool
case object True extends Bool
case object False extends Bool

def t = True
def f = False

def id(x : Bool) : Bool = x

def not(x : Bool) : Bool = x match {
    case True => False
    case False => True
}

def and(x : Bool, y : Bool) : Bool = (x,y) match {
    case (True, True) => True
    case _ => False
}

def or(x : Bool, y : Bool) : Bool = (x, y) match {
    case (False, False) => False
    case _              => True
}

def xor(x : Bool, y : Bool) : Bool = (x, y) match{
    case (True, False) => True
    case (False, True) => True
    case _ => False
}

def nand(x : Bool, y : Bool) : Bool = not(and(x,y))

def bool_eq(x : Bool, y : Bool) : Bool = not(xor(x,y))

// Lists
sealed trait List[+A]
case object Empty extends List[Nothing]
case class Cons[A](x : A, xs : List[A]) extends List[A]

def map[A, B](f : (A => B), xs : List[A]) : List[B] = xs match {
    case Empty       => Empty
    case Cons(x, xs) => Cons(f(x), map(f, xs))
}

def fold[A, B](f : ((A, B) => B), acc : B, xs : List[A]) : B = xs match {
    case Empty       => acc
    case Cons(x, xs) => fold(f, f(x, acc), xs)
}

def filter[A](p : (A => Bool), xs : List[A]) : List[A] = xs match {
    case Empty       => Empty
    case Cons(x, xs) => p(x) match {
        case True  => Cons(x, filter(p, xs))
        case False => filter(p, xs)
    }
}

// Maybe
sealed trait Maybe[+A]
case object Nothing extends Maybe[Nothing]
case class Just[A](fromJust : A) extends Maybe[A]

def map_maybe[A, B](f : (A => B), m : Maybe[A]) : Maybe[B] = m match{
    case Nothing => Nothing
    case Just(x) => Just(f(x))
}

// Maps
sealed trait Dictionary[+A, +B] 
case object EmptyDict extends Dictionary[Nothing, Nothing]
case class Entry[A,B](key : A, value : B, xs : Dictionary[A, B]) extends Dictionary[A, B]

def isIn[A,B](eq : ((A,A) => Bool), dict : Dictionary[A,B], k : A ) : Bool = dict match {
    case EmptyDict        => False
    case Entry(k1, v, xs) => eq(k,k1) match {
        case True  => True
        case False => isIn(eq, dict, k)
    }
}

def lookup[A,B](eq : (A,A) => Bool, dict : Dictionary[A,B], k : A) : Maybe[B] = dict match {
    case EmptyDict        => Nothing
    case Entry(k1, v, xs) => eq(k,k1) match {
        case True  => Just(v)
        case False => lookup(eq, dict, k)
    }
}

// Strings
def string_eq(s1 : String, s2 : String) : Bool = if (s1 == s2) True else False

// Products
sealed trait Pair[+A, +B]
case class MkPair[A, B](fst : A, snd : B) extends Pair[A, B]

// Sums
sealed trait Either[+A, +B]
case class Left[A, B](left : A) extends Either[A, B]
case class Right[A, B](right : B) extends Either[A, B]
```

Output:
```
defined trait Nat
defined object Zero
defined class Succ
one: Succ = Succ(Zero)
two: Succ = Succ(Succ(Zero))
three: Succ = Succ(Succ(Succ(Zero)))
four: Succ = Succ(Succ(Succ(Succ(Zero))))
five: Succ = Succ(Succ(Succ(Succ(Succ(Zero)))))
six: Succ = Succ(Succ(Succ(Succ(Succ(Succ(Zero))))))
seven: Succ = Succ(Succ(Succ(Succ(Succ(Succ(Succ(Zero)))))))
eight: Succ = Succ(Succ(Succ(Succ(Succ(Succ(Succ(Succ(Zero))))))))
nine: Succ = Succ(Succ(Succ(Succ(Succ(Succ(Succ(Succ(Succ(Zero)))))))))
ten: Succ = Succ(Succ(Succ(Succ(Succ(Succ(Succ(Succ(Succ(Succ(Zero))))))))))
nat_plus: (x: Nat, y: Nat)Nat
nat_to_int: (x: Nat)Int
nat_to_str: (x: Nat)String
print_nat: (x: Nat)Unit
nat_mult: (x: Nat, y: Nat)Nat
nat_minus: (x: Nat, y: Nat)Nat
nat_pow: (x: Nat, y: Nat)Nat
nat_lte: (x: Nat, y: Nat)Bool
nat_eq...
```

# Homework 6

In this assignment we will write the first full interpreter for Lettuce. Remember that lettuce is a functional language with let bindings. It is very similiar to a language called ML. Here we want to implement the interpreter for the language from the ground up. We will use most of the existing standard library we have developed while we write this, especially the equality functions and Dictionary data type from the last homework.

## Part 1 - The Values

Recall that we define interpreters as functions that take in expressions in the form of abstract syntax and give a value as output. In symbols:

$$
eval : Expr \rightarrow Value
$$

It follows that we will need to create a datatype to represent the values that may be computed by lettuce. Below is the grammar that defines lettuce values. Implement this as a `sealed trait` in Scala below:

$$
\begin{align}
\textbf{Value} ::=&\ NumVal\ \mathbb{Z}\\
\mid&\ BinVal\ \mathbb{B}\\
\mid&\ Error
\end{align}
$$

<!-- grade_id: values-adt, points: 0, grade: false, solution: true -->
```scala
sealed trait Value

case object Error extends Value
case class NumVal(x: Integer) extends Value
case class BinVal(b: Bool) extends Value
```

Output:
```
defined trait Value
defined object Error
defined class NumVal
defined class BinVal
warning: previously defined object Value is not a companion to trait Value.
Companions must be defined together; you may wish to use :paste mode for this.
```

If your definition was correct then the values in the cell below should compile.

<!-- grade_id: values-adt-test, points: 5, grade: true, solution: false -->
```scala
val v1 = NumVal(Positive(Succ(Succ(Zero))))
val v2 = BinVal(False)
val v3 = Error
```

Output:
```
Error
```

## Part 2 - The Expressions

There are many possible expressions in the Lettuce language. We also use an abstract syntax to represent this. Here is the grammar to remind you of all the syntactic elements in Lettuce:

$$\begin{array}{rcll}
\mathbf{Expr} & ::= & Const(\mathbb{Z}) \\
 & | & Bin(\mathbb{B}) \\
 & | & Ident(\mathbf{String}) \\
 & | & Plus(\mathbf{Expr}, \mathbf{Expr}) \\
 & | & Minus(\mathbf{Expr}, \mathbf{Expr}) \\
 & | & Mult (\mathbf{Expr}, \mathbf{Expr}) \\
 & | & Pow (\mathbf{Expr}, \mathbf{Expr}) \\
 & | & Neg (\mathbf{Expr}) \\
 & | & Eq (\mathbf{Expr}, \mathbf{Expr}) \\
 & | & And ( \mathbf{Expr}, \mathbf{Expr} ) \\
 & | & Or ( \mathbf{Expr}, \mathbf{Expr} ) \\
 & | & IfThenElse(\mathbf{Expr}, \mathbf{Expr}, \mathbf{Expr}) \\
 & | & Let( \mathbf{String}, \mathbf{Expr}, \mathbf{Expr}) \\
\end{array}$$

Now define a `sealed trait` for expressions in Lettuce:

<!-- grade_id: expr, points: 0, grade: false, solution: true -->
```scala
sealed trait Expr

case class Const(x: Integer) extends Expr
case class Bin(b: Bool) extends Expr
case class Ident(s: String) extends Expr
case class Plus(e1: Expr, e2: Expr) extends Expr
case class Minus(e1: Expr, e2: Expr) extends Expr
case class Mult(e1: Expr, e2: Expr) extends Expr
case class Pow(e1: Expr, e2: Expr) extends Expr
case class Neg(e: Expr) extends Expr
case class Eq(e1: Expr, e2: Expr) extends Expr
case class And(e1: Expr, e2: Expr) extends Expr
case class Or(e1: Expr, e2: Expr) extends Expr
case class IfThenElse(e1: Expr, e2: Expr, e3: Expr) extends Expr
case class Let(s: String, e1: Expr, e2: Expr) extends Expr
```

Output:
```
defined trait Expr
defined class Const
defined class Bin
defined class Ident
defined class Plus
defined class Minus
defined class Mult
defined class Pow
defined class Neg
defined class Eq
defined class And
defined class Or
defined class IfThenElse
defined class Let
```

If you defined `Expr` properly then the following expressions should compile:

<!-- grade_id: expr-test, points: 5, grade: true, solution: false -->
```scala
val ex1 = Const(Positive(five))
val ex2 = Const(Negative(three))
val ex3 = Bin(True)
val ex4 = Ident("x")
val ex5 = Plus(ex1, ex2)
val ex6 = Let("x", ex5, Mult(ex4, ex4))
```

Output:
```
Let(x,Plus(Const(Positive(Succ(Succ(Succ(Succ(Succ(Zero))))))),Const(Negative(Succ(Succ(Succ(Zero)))))),Mult(Ident(x),Ident(x)))
```

## Part 3 - The Interpreter 

We now need to define the interpreter for this language. It is a function of the form:

$$
eval : Expr \rightarrow Value
$$

Except, now that we have an execution environment($\sigma$), we need some way to include $\sigma$ in this function. This will be as an additional argument:

$$
eval : \sigma \rightarrow Expr \rightarrow Value
$$

We will let $\sigma$ be a `Dictionary` as we defined in last week's homework. To add something to a dicitonary use the `Entry` constructor to add on a new element. This should be very similiar to the way `Cons` works on lists.

Below we have given you the skeleton of the function `eval`. Fill in each case for the interpreter so that it can interpret all possible Lettuce expressions. The inference rules we covered in the classroom should be very helpful with this.

Hint: It may be helpful to define some helper functions for the inference rules that have similar forms. These are the rules like binary operations for arithmetic, bool, etc. Just like we lumped some of these rules together into a single rule, we can do the same for our interpreter in the form of auxillary functions. The bonus `eval_bin` function from Homework 4 may be a good place for inspiration.

Hint: Be ready to write some nested case matches. Most of your cases should have two nestings. One for the Sytnax element and another for evaluating the expressions it may contain. This is not true for all expressions but most will follow this pattern to some extent. Note that `Ident` has a different kind of case match because it needs to interact with the environment(a Dictionary)

<!-- grade_id: interpreter, points: 0, grade: false, solution: true -->
```scala
def eval(env: Dictionary[String, Value], expr: Expr): Value = expr match {
//     case Const(x: Integer) => x match {
//         case Positive(x) => NumVal(Positive(x))
//         case Negative(x) => NumVal(Negative(x))
//     }
    x match {
        
    }
    case Bin(b: Bool) => BinVal(b)
    case Ident(s: String) => Error
    case Plus(e1: Expr, e2: Expr) => Error
    case Minus(e1: Expr, e2: Expr) => Error
    case Mult(e1: Expr, e2: Expr) => Error
    case Pow(e1: Expr, e2: Expr) => Error
    case Neg(e: Expr) => Error
    case Eq(e1: Expr, e2: Expr) => Error
    case And(e1: Expr, e2: Expr) => Error
    case Or(e1: Expr, e2: Expr) => Error
    case IfThenElse(e1: Expr, e2: Expr, e3: Expr) => Error
    case Let(s: String, e1: Expr, e2: Expr) => Error
    case _ => Error
}
```

Output:
```
eval: (env: Dictionary[String,Value], expr: Expr)Value
```

If this was defined correctly then the definitions above should evaluate:

<!-- grade_id: interpreter-tests, points: 20, grade: true, solution: false -->
```scala
//val ex1 = Const(Positive(five))
//val ex2 = Const(Negative(three))
//val ex3 = Bin(True)
//val ex4 = Ident("x")
//val ex5 = Plus(ex1, ex2)
//val ex6 = Let("x", ex5, Mult(ex4, ex4))


assert(eval(EmptyDict, ex3) equals BinVal(True))
assert(eval(EmptyDict, ex4) equals Error)
assert(eval(EmptyDict, ex5) equals NumVal(Positive(two)))
assert(eval(EmptyDict, ex6) equals NumVal(Positive(four)))
```

Output:
```
java.lang.AssertionError: assertion failed
  at scala.Predef$.assert(Predef.scala:156)
  ... 44 elided
```

```scala

```